```python
# Initialize Otter
import otter
grader = otter.Notebook("hw1.ipynb")
```

# HW1: Introduction to Pricing

Cost-Plus Pricing, Competitive Parity Pricing, Value-Based Pricing

```python
import polars as pl
import numpy as np
from plotnine import ggplot, aes, geom_histogram, geom_line, geom_point, labs, theme_minimal, theme
```

### Part 1: Cost-Plus Pricing

You are advising a new restaurant in La Jolla that specializes in Salads and Soups. Following the convention in the restaurant industry, the owner wants to follow a cost-plus pricing strategy to set menu prices. Their target is to maintain a 50% margin over cost on each item sold. Based on his experience with other restaurants, they expect to sell 100,000 salads and 50,000 soups in the first year of operation.

Their material costs are $5.00 per salad and $3.00 per cup of soup. Rent is $300,000 per year. 

You goal is to set the price for salads and soups.

```python
material_cost_salad = 5.00
material_cost_soup = 3.00
rent = 300000
expected_sales_salad = 100000
expected_sales_soup = 50000
target_margin = 0.50
```

**Question 1:** You need to allocate the fixed costs (rent) to each menu item before calculating the price. Assume the owner allocates the fixed costs based on the expected sales volume of each item. What is the allocated fixed cost per salad and per soup? What is the price for each item using the cost-plus pricing method? 

```python
cost_salad_q1 = ...
cost_soup_q1 = ...

price_salad_q1 = ...
price_soup_q1 = ...

print("When attributing fixed costs on the basis of expected sales volume for each item:")
print(f"The unit cost for salad is ${cost_salad_q1:.2f}, and the price is ${price_salad_q1:.2f}.")
print(f"The unit cost for soup is ${cost_soup_q1:.2f}, and the price is ${price_soup_q1:.2f}.")
```

```python
grader.check("q1")
```

**Question 2:** Now assume the owner allocates the fixed costs based on the expected total material costs of each item. What is the allocated fixed cost per salad and per soup? What is the price for each item using the cost-plus pricing method? 

```python
total_material_cost_salad = material_cost_salad * expected_sales_salad
total_material_cost_soup = material_cost_soup * expected_sales_soup

cost_salad_q2 = ...
cost_soup_q2 = ...

price_salad_q2 = ...
price_soup_q2 = ...
print("When attributing fixed costs on the basis of total material costs for each item:")
print(f"The new unit cost for salad is ${cost_salad_q2:.2f}, and the new price is ${price_salad_q2:.2f}.")
print(f"The new unit cost for soup is ${cost_soup_q2:.2f}, and the new price is ${price_soup_q2:.2f}.")
```

```python
grader.check("q2")
```

**Question 3:** Now suppose you have run a survey and found that there are essentially two types of local customers: 
- UCSD Students: who *all* value soup at $6 and salad at $10.50
- Tourists: who *all* value soup at $8 and salad at $12.50.

UCSD students are willing to consume 50,000 salads and 25,000 soups. Tourists are also willing to consume 50,000 salads and 25,000 soups.

Based on your answers to questions 1 and 2, what are:
- Total profits under cost-plus pricing using expected sales volume to allocate fixed costs?
- Total profits under cost-plus pricing using expected material costs to allocate fixed costs?

```python
student_wtp_salad = 10.50
student_wtp_soup = 6.00
tourist_wtp_salad = 12.50
tourist_wtp_soup = 8.00

student_demand_salad = 50000
student_demand_soup = 25000
tourist_demand_salad = 50000
tourist_demand_soup = 25000
```

```python
# Question 1 pricing: using expected sales volume to allocate fixed costs
demand_salad_q1 = ...
demand_soup_q1 = ...
profit_q1 = ...

print("Using cost-plus pricing with expected sales volume to allocate fixed costs:")
print(f"Salad price: ${price_salad_q1:.2f}, Soup price: ${price_soup_q1:.2f}")
print(f"Total demand for salad is {demand_salad_q1:,}, total demand for soup is {demand_soup_q1:,} and profit is ${profit_q1:,.2f}.")

# Question 2 pricing: using total material costs to allocate fixed costs
demand_salad_q2 = ...
demand_soup_q2 = ...
profit_q2 = ...

print("Using cost-plus pricing with total material costs to allocate fixed costs:")
print(f"Salad price: ${price_salad_q2:.2f}, Soup price: ${price_soup_q2:.2f}")
print(f"Total demand for salad is {demand_salad_q2:,}, total demand for soup is {demand_soup_q2:,} and profit is ${profit_q2:,.2f}.")

```

```python
grader.check("q3")
```

**Question 4 (multiple choice):** Based on the previous question, which cost-plus pricing method is better?

Your answer can be one of the following options:
- A) Using **expected sales volume** to allocate fixed costs results in higher profits
- B) Using **expected material costs** to allocate fixed costs results in higher profits

Write your answer by assigning a string variable `q4_answer` with the corresponding letter (e.g., `q4_answer = "X"`).

```python
q4_answer = ...
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

**Question 5 (open ended):** What other pricing framework can explain why one of the cost-plus pricing methods yields a higher profit?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Part 2: Value-Based Pricing

While working for HP's Commercial Printing division, you have to price a brand new commercial photo
printer (Z53123) for the Southern California Area. This new printer represents a significant change from
the current market offerings (assume there is no competition!). You identified 1,000 possible print shops
that could consider purchasing this printer. With in-depth surveys, market research and data analysis1,
your team has estimated the willingness to pay for each consumer.

The raw data are available in a CSV `hw1_data.csv`

The accounting team at HP estimates that the variable cost to produce each machine is $6,000 and the
fixed cost for renting a manufacturing facility in Southern California is $1 million. 

You are considering prices between $5,000 and $15,000 in increments of $1,000.

```python
df_wtp = pl.read_csv("hw1_data.csv")

cost_per_machine = 6000
fixed_cost = 1e6

# Plotting function - Histogram
def plot_histogram(data, column, bins=10, title="", xlabel="", ylabel="Frequency"):
    # If bins is an array (bin edges), use breaks parameter; otherwise use bins
    if isinstance(bins, np.ndarray):
        p = (
            ggplot(data, aes(x=column))
            + geom_histogram(breaks=bins, color='white')
            + labs(title=title, x=xlabel, y=ylabel)
            + theme_minimal()
            + theme(figure_size=(10, 6))
        )
    else:
        p = (
            ggplot(data, aes(x=column))
            + geom_histogram(bins=bins, color='white')
            + labs(title=title, x=xlabel, y=ylabel)
            + theme_minimal()
            + theme(figure_size=(10, 6))
        )
    return p

# Plotting function - Line Chart
def plot_line_chart(data, x_column, y_column, title="", xlabel="", ylabel=""):
    p = (
        ggplot(data, aes(x=x_column, y=y_column))
        + geom_line()
        + geom_point()
        + labs(title=title, x=xlabel, y=ylabel)
        + theme_minimal()
        + theme(figure_size=(10, 6))
    )
    return p
```

The histogram of willingness to pay values is shown below.

```python
plot_histogram(df_wtp, column="WTP", bins = np.linspace(5000,15000,11), title="Histogram of Willingness to Pay", xlabel="Willingness to Pay ($)")
```

**Question 6:** Compute and plot the demand curve in incremental prices of $1,000 between $5,000 and $15,000. The autograder will only check that your demand values are correct.

```python
price_points = np.arange(5000, 16000, 1000)

# Define demand function
demand = ...

demand_curve = np.array([demand(price) for price in price_points])

demand_df = pl.DataFrame({
    "Price": price_points,
    "Demand": demand_curve
})

plot_line_chart(demand_df, x_column="Price", y_column="Demand", title="Demand Curve", xlabel="Price ($)", ylabel="Demand (Number of Customers)")
```

```python
grader.check("q6")
```

**Question 7:** Compute and plot the profit curve for prices between $5,000 and $15,000 in increments of $1,000. Determine the profit-maximizing price.

The autograder will only check that your profit values and profit-maximizing price are correct.

```python
profit_curve = ...

profit_df = pl.DataFrame({
    "Price": price_points,
    "Profit": profit_curve
})

plot_line_chart(profit_df, x_column="Price", y_column="Profit", title="Profit Curve", xlabel="Price ($)", ylabel="Profit ($)")
```

```python
profit_max_price = ...

print(f"Profit-maximizing price: ${profit_max_price:,}")
print(f"Maximum profit at this price: ${profit_curve[np.argmax(profit_curve)]:,.2f}")
print(f"Demand at profit-maximizing price: {demand_curve[np.argmax(profit_curve)]} customers")
```

```python
grader.check("q7")
```

**Question 8:** Now suppose capacity were limited to producing 500 machines. Re-compute the demand curve and profit curve under this capacity constraint and plot the profit curve. 

What is the new profit-maximizing price?

How much would you be willing to pay to increase capacity from 500 to 1,000 machines (so that you can sell to all willing customers at the profit-maximizing price as in Question 6)?

The autograder will only check that your profit curve, profit-maximizing price, and willingness to pay for capacity increase are correct.

```python
capacity_limit = 500

demand_curve_capped = ...

profit_curve_capped = ...

profit_capped_df = pl.DataFrame({
    "Price": price_points,
    "Profit": profit_curve_capped
})

plot_line_chart(profit_capped_df, x_column="Price", y_column="Profit", title="Profit Curve under Capacity Constraint", xlabel="Price ($)", ylabel="Profit ($)")
```

```python
profit_max_price_capped = ...
```

```python
wtp_added_capacity = ...

print(f"With capacity limit of {capacity_limit} machines:")
print(f"Profit-maximizing price under constraint: ${profit_max_price_capped:,}")
print(f"Maximum profit under constraint: ${profit_curve_capped[np.argmax(profit_curve_capped)]:,.2f}")
print(f"Willingness to pay for added capacity: ${wtp_added_capacity:,.2f}")
```

```python
grader.check("q8")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you have the zip file, upload the **entire** zip file to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 15)
```
>>> assert isinstance(cost_salad_q1, (int, float)), 'cost_salad_q1 should be a single number'
>>> assert isinstance(cost_soup_q1, (int, float)), 'cost_soup_q1 should be a single number'
>>> assert isinstance(price_salad_q1, (int, float)), 'price_salad_q1 should be a single number'
>>> assert isinstance(price_soup_q1, (int, float)), 'price_soup_q1 should be a single number'
>>> True
True
```


Tests for q2 (points: 15)
```
>>> assert isinstance(cost_salad_q2, float), 'cost_salad_q2 should be a single float value'
>>> assert isinstance(cost_soup_q2, float), 'cost_soup_q2 should be a single float value'
>>> assert isinstance(price_salad_q2, float), 'price_salad_q2 should be a single float value'
>>> assert isinstance(price_soup_q2, float), 'price_soup_q2 should be a single float value'
>>> True
True
```


Tests for q3 (points: 15)
```
>>> assert isinstance(demand_salad_q1, (int, float, np.integer)), 'demand_salad_q1 should be a number'
>>> assert isinstance(demand_soup_q1, (int, float, np.integer)), 'demand_soup_q1 should be a number'
>>> assert isinstance(profit_q1, (int, float, np.integer, np.floating)), 'profit_q1 should be a number'
>>> assert isinstance(demand_salad_q2, (int, float, np.integer)), 'demand_salad_q2 should be a number'
>>> assert isinstance(demand_soup_q2, (int, float, np.integer)), 'demand_soup_q2 should be a number'
>>> assert isinstance(profit_q2, (int, float, np.integer, np.floating)), 'profit_q2 should be a number'
>>> True
True
```


Tests for q4 (points: 5)
```
>>> assert q4_answer in ['A', 'B'], "q4_answer should be either 'A' or 'B'"
>>> True
True
```


Tests for q6 (points: 10)
```
>>> assert isinstance(demand_curve, np.ndarray), 'demand_curve should be a numpy array'
>>> assert len(demand_curve) == 11, 'demand_curve should have 11 values (prices from $5,000 to $15,000)'
>>> assert all((d >= 0 for d in demand_curve)), 'All demand values should be non-negative'
>>> True
True
```


Tests for q7 (points: 15)
```
>>> assert isinstance(profit_curve, np.ndarray), 'profit_curve should be a numpy array'
>>> assert len(profit_curve) == 11, 'profit_curve should have 11 values'
>>> assert isinstance(profit_max_price, (int, np.integer)), 'profit_max_price should be an integer'
>>> assert 5000 <= profit_max_price <= 15000, 'profit_max_price should be between $5,000 and $15,000'
>>> True
True
```


Tests for q8 (points: 15)
```
>>> assert isinstance(profit_curve_capped, np.ndarray), 'profit_curve_capped should be a numpy array'
>>> assert len(profit_curve_capped) == 11, 'profit_curve_capped should have 11 values'
>>> assert isinstance(profit_max_price_capped, (int, np.integer)), 'profit_max_price_capped should be an integer'
>>> assert 5000 <= profit_max_price_capped <= 15000, 'profit_max_price_capped should be between $5,000 and $15,000'
>>> assert isinstance(wtp_added_capacity, (int, float, np.integer, np.floating)), 'wtp_added_capacity should be a number'
>>> True
True
```
