```python
# Initialize Otter
import otter
grader = otter.Notebook("hw3.ipynb")
```

# HW3: Choice-Based Conjoint Analysis

A sports car manufacturer conducted a conjoint survey where 200 respondents each answered 10 choice questions. In each question, respondents chose between 3 sports car configurations described by:

- **Seats**: 2, 4, or 5
- **Transmission**: Manual or Automatic
- **Convertible**: Yes or No
- **Price**: $30K, $35K, or $40K

The file `hw3.csv` contains the survey data, with additional dummy variables for 4-seat, 5-seat, and automatic transmission options.

```python
import polars as pl
import numpy as np
import matplotlib.pyplot as plt
from xlogit import MultinomialLogit
```

```python
df = pl.read_csv("hw3_data.csv")
df.head(10)
```

## Part 1: Model Estimation

We model choice probability using the multinomial logit:

$$P(\text{choice} = j) = \frac{e^{V_j}}{\sum_{k} e^{V_k}}$$

where $V_j = \beta_{seat4} \cdot \mathbb{1}[seat=4] + \beta_{seat5} \cdot \mathbb{1}[seat=5] + \beta_{auto} \cdot \mathbb{1}[auto] + \beta_{conv} \cdot \mathbb{1}[convertible] + \beta_{price} \cdot price$

To estimate the model, we will use the `xlogit` python package. The package requires us to provide a numpy array of covariates `X`, a numpy vector of choices (either 0 or 1) `choice`, a list of varnames, a numpy vector of alternative IDs and a numpy vector of choice IDs (i.e. for each respondent and survey question).

```python
# Prepare data for estimation
varnames = ["seat4", "seat5", "auto", "convert", "price"]
X = df.select(varnames).to_numpy()
choice = df["choice"].to_numpy()
choice_id = df["choice_id"].to_numpy()
alt_id = df["alt"].to_numpy()
```

```python
model = MultinomialLogit()
model.fit(X=X, y=choice, varnames=varnames, alts=alt_id, ids=choice_id, fit_intercept=False)

# Save the coefficients into a dictionary for later use
coefs = dict(zip(model.coeff_names, model.coeff_))

# Print the model summary
model.summary()
```

**Question 1 (5 pts):** Which of the estimated coefficients are statistically significant at the 1% level?

Mark true or false on each of the following options.

```python
q1_answer = {
    ...
    ,
    ...
    ,
    ...
    ,
    ...
    ,
    ...
}
q1_answer
```

## Part 2: Willingness to Pay

Now calculate the willingness to pay for each of the following features:
- 5 seats instead of 2 seats
- Automatic transmission instead of manual
- Convertible instead of non-convertible

Recall that willingness to pay involves the trade-off between how strongly you *value* a feature and how strongly you *dislike* a higher price.

**Question 2 (10 pts):** Calculate the willingness to pay for each feature (in $K). Store the results in `wtp_seat5`, `wtp_auto`, and `wtp_convertible`.

```python
# You can access the coefficients from the regression 

wtp_seat5 = ...
wtp_auto = ...
wtp_convertible = ...

print(f"WTP for convertible: ${wtp_convertible:.2f}K")
print(f"WTP for automatic transmission: ${wtp_auto:.2f}K")
print(f"WTP for 5 seats: ${wtp_seat5:.2f}K")
```

```python
grader.check("q2")
```

## Part 3: Market Share and Demand

The power of this model is we can now flexibly predict market shares and demand for different product configurations and prices.

Recall from class that the market share for product $j$ is:

$$Share_j = \frac{e^{V_j}}{\sum_{k} e^{V_k}}$$

(Note that the outside option is not included in this formula since we are only interested in market shares among the offered products, which is consistent with our survey design.)

For convenience, the function below computes market shares for a given set of product attributes and prices.

```python
def utility(attributes, price, coefs):
    """Calculate utility for a product configuration."""
    assert isinstance(attributes, dict), "Attributes must be provided as a dictionary."
    assert attributes.keys() == {"seat", "trans", "convertible"}, "Attributes dictionary must contain 'seat', 'trans', and 'convertible' keys."
    assert attributes["seat"] in [2, 4, 5], "Seat attribute must be 2, 4, or 5."
    assert attributes["trans"] in ["auto", "manual"], "Transmission attribute must be 'auto' or 'manual'."
    assert attributes["convertible"] in ["yes", "no"], "Convertible attribute must be 'yes' or 'no'."    
    v = coefs["price"] * price
    if attributes["seat"] == 4: v += coefs["seat4"]
    elif attributes["seat"] == 5: v += coefs["seat5"]
    if attributes["trans"] == "auto": v += coefs["auto"]
    if attributes["convertible"] == "yes": v += coefs["convert"]
    return v

def market_shares(attributes, prices, coefs):
    """Calculate market shares with outside option."""
    exp_v = [np.exp(utility(attr, price, coefs=coefs)) for attr, price in zip(attributes, prices)]
    denom = sum(exp_v)
    return [e / denom for e in exp_v]
```

For example, we can compute market shares for the following three products:
1. 4-seat, automatic, non-convertible, $35K
2. 5-seat, automatic, convertible, $40K
3. 2-seat, manual, non-convertible, $32K


```python
attributes = [
    {"seat": 4, "trans": "auto", "convertible": "no"},
    {"seat": 5, "trans": "auto", "convertible": "yes"},
    {"seat": 2, "trans": "manual", "convertible": "no"},
]

prices = [35, 40, 32]

shares = market_shares(attributes, prices, coefs=coefs)
for i, share in enumerate(shares):
    print(f"Product {i+1} market share: {share*100:.2f}%")
```

## Part 4: Optimal Pricing

You are launching the **Thunderbolt**, a new sports car:
- 4 seats, automatic, non-convertible
- Marginal cost: $25K

The competitor **Speedster** is already in the market:
- 2 seats, manual, convertible
- Price: $35K (fixed)

Total market size: 10,000 potential buyers.



```python
thunderbolt_attrs = {"seat": 4, "trans": "auto", "convertible": "no"}
speedster_attrs = {"seat": 2, "trans": "manual", "convertible": "yes"}
speedster_price = 35 # In thousands of dollars

marginal_cost = 25 # In thousands of dollars
market_size = 10_000
thunderbolt_prices = np.arange(26, 51)
```

**Question 3 (10 pts):** Compute the demand for the Thunderbolt at each price point from $26K to $50K in increments of $1K. Remember that the price variable we have been using is already in thousands of dollars.

Store the result in `thunderbolt_demand` as a numpy array.

You can use the `market_shares` function defined above to compute the market share for the Thunderbolt at each price point, and then multiply by the total market size to get demand. Remember, the function takes in a list of attribute dictionaries and a list of prices. So you simply need to substitute in the Thunderbolt and Speedster attributes, the Speedster price, and try different candidate prices for the Thunderbolt.

```python
def demand(thunderbolt_prices, thunderbolt_attrs, speedster_attrs, speedster_price, coefs=coefs):
    demand = np.zeros(thunderbolt_prices.shape)
    for i, price in enumerate(thunderbolt_prices):
        attributes = [
            thunderbolt_attrs,
            speedster_attrs
        ]
        price_list = ...
        shares = ...
        thunderbolt_share = shares[0]
        demand[i] = thunderbolt_share * market_size
    return demand

thunderbolt_demand = ...
```

```python
# Plot the demand curve
plt.figure(figsize=(10, 6))
plt.plot(thunderbolt_prices, thunderbolt_demand, marker='o')
plt.title("Demand Curve for Thunderbolt")
plt.xlabel("Price ($K)")
plt.ylabel("Demand (units)")
plt.grid()
plt.show()
```

```python
grader.check("q3")
```

**Question 4 (10 pts):** Compute the profit at each price point. Store in `thunderbolt_profit` as a numpy array. Then find the profit-maximizing price and store it in `optimal_price`.

```python
thunderbolt_profit = ...
optimal_price = ...
```

```python
# Plot the profit curve and the optimal price point
plt.figure(figsize=(10, 6))
plt.plot(thunderbolt_prices, thunderbolt_profit/1000, marker='o', label='Profit')
plt.axvline(optimal_price, color='r', linestyle='--', label='Optimal Price')
plt.title("Profit Curve for Thunderbolt")
plt.xlabel("Price ($K)")
plt.ylabel("Profit ($M)")
plt.legend()
plt.grid()
plt.show()
```

```python
grader.check("q4")
```

## Part 5: Product Design Decision

The head of engineering at the company has proposed making the Thunderbolt a convertible. This would increase the marginal cost and require an upfront investment, but could also increase demand.

Here are the details:
- Additional marginal cost: $700 per unit
- Upfront investment: $8M for retooling and R&D
- Planning horizon: 3 years at 10,000 buyers/year

As head of pricing, you are asked for your input on this business decision.

```python
thunderbolt_conv_attrs = {"seat": 4, "trans": "auto", "convertible": "yes"}
mc_convertible = 25.7  # $25K + $700
upfront = 8000  # $8M in $K
years = 3
```

**Question 5 (10 pts):** Repeat the demand calculation for the convertible Thunderbolt at each price point from $26K to $50K. Store in `thunderbolt_convertible_demand`.

```python
# You can re-use the demand function from above
thunderbolt_convertible_demand = ...
```

```python
# Plot the demand curve for the convertible Thunderbolt against the non-convertible version
plt.figure(figsize=(10, 6))
plt.plot(thunderbolt_prices, thunderbolt_demand, marker='o', label='Non-convertible')
plt.plot(thunderbolt_prices, thunderbolt_convertible_demand, marker='o', label='Convertible')
plt.title("Demand Curve for Thunderbolt: Convertible vs Non-convertible")
plt.xlabel("Price ($K)")
plt.ylabel("Demand (units)")
plt.legend()
plt.grid()
plt.show()
```

```python
grader.check("q5")
```

**Question 6 (10 pts):** Compute the profit curve and profit-maximizing price for the convertible Thunderbolt. Store in `thunderbolt_conv_profit` and `optimal_conv_price`.

```python
# Compute the profit curve and the optimal price for the convertible Thunderbolt
thunderbolt_conv_profit = ...
optimal_conv_price = ...
```

```python
# Plot the profit curves for the convertible and non-convertible Thunderbolt before upfront cost
plt.figure(figsize=(10, 6))
plt.plot(thunderbolt_prices, thunderbolt_profit/1000, marker='o', label='Non-convertible')
plt.plot(thunderbolt_prices, thunderbolt_conv_profit/1000, marker='o', label='Convertible')
plt.axvline(optimal_price, color='tab:blue', linestyle='--', label='Optimal Non-convertible Price')
plt.axvline(optimal_conv_price, color='tab:orange', linestyle='--', label='Optimal Convertible Price')
plt.title("Profit Curves for Thunderbolt: Convertible vs Non-convertible (before upfront cost)")
plt.xlabel("Price (in $K)")
plt.ylabel("Profit ($M)")
plt.legend()
plt.grid()
plt.show()
```

```python
grader.check("q6")
```

**Question 7 (5 pts):** Compute the difference in 3-year total profit (after upfront investment) between the convertible and non-convertible versions. Is the convertible upgrade worth the investment? Set `convertible_worth_it` to `True` or `False`.

```python
profit_diff = ...
convertible_worth_it = ...
```

```python
grader.check("q7")
```

## Part 6: Segment-Specific Analysis


The survey classified respondents into segments: "basic", "fun", and "racer". The executive team has decided to fully target only the "fun" segment with the Thunderbolt. These represent approximately 2,550 potential buyers in the market.

```python
# Check segment distribution
df.filter(pl.col("alt") == 1).group_by("segment").len()
```

```python
market_size_fun = market_size * (df['segment'] == 'fun').mean()
```

**Question 8 (10 pts):** Re-estimate the model using only the "fun" segment. Store the coefficients in `coefs_fun` and calculate the WTP for convertible (`wtp_conv_fun`).

```python
df_fun = df.filter(pl.col("segment") == "fun")

model_fun = MultinomialLogit()
model_fun.fit(
    X=df_fun.select(varnames).to_numpy(),
    y=df_fun["choice"].to_numpy(),
    varnames=varnames,
    alts=df_fun["alt"].to_numpy(),
    ids=df_fun["choice_id"].to_numpy(),
    fit_intercept=False
)
model_fun.summary()

coefs_fun = ...
wtp_conv_fun = ...

print(f"\nWTP for convertible (fun segment): ${wtp_conv_fun:.2f}K")
print(f"WTP for convertible (full sample): ${wtp_convertible:.2f}K")
```

```python
grader.check("q8")
```

**Question 9 (10 pts):** Using the "fun" segment coefficients, find the optimal prices for both non-convertible and convertible Thunderbolts. Store results in:
- `optimal_price_fun`: optimal price for non-convertible (using fun coefs)
- `optimal_conv_price_fun`: optimal price for convertible (using fun coefs)

```python
# Compute demand, profit, and optimal price for the fun segment
thunderbolt_demand_fun = ...
profits_fun = ...
optimal_price_fun = ...

# Compute demand, profit, and optimal price for the fun segment convertible
thunderbolt_convertible_demand_fun = ...
thunderbolt_conv_profit_fun = ...
optimal_conv_price_fun = ...
```

```python
# Plot the new profit curves under the fun segment demand for the convertible and non-convertible Thunderbolt before upfront cost
plt.figure(figsize=(10, 6))
plt.plot(thunderbolt_prices, profits_fun/1000, marker='o', label='Non-convertible')
plt.plot(thunderbolt_prices, thunderbolt_conv_profit_fun/1000, marker='o', label='Convertible')
plt.axvline(optimal_price_fun, color='tab:blue', linestyle='--', label='Optimal Non-convertible Price')
plt.axvline(optimal_conv_price_fun, color='tab:orange', linestyle='--', label='Optimal Convertible Price')
plt.title("Within Fun Segment Profit: Convertible vs Non-convertible (before upfront cost)")
plt.xlabel("Price (in $K)")
plt.ylabel("Profit ($M)")
plt.legend()
plt.grid()
plt.show()
```

```python
grader.check("q9")
```

**Question 10 (5 pts):** Compute the difference in 3-year total profit (after upfront investment) between the convertible and non-convertible versions with the **fun segment demand**. Is the convertible upgrade worth the investment? Set `convertible_worth_it_fun_segment` to `True` or `False`.

```python
profit_diff_fun = ...
convertible_worth_it_fun_segment = ...
```

```python
grader.check("q10")
```

<!-- BEGIN QUESTION -->

**Question 11 (10 pts, Open-Ended):** The convertible upgrade was not worth it using the full sample, but became profitable when targeting the "fun" segment. Explain why accounting for preference heterogeneity across customer segments is important for pricing and product design decisions.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 12 (5 pts, Open-Ended):** What are some limitations of using the above choice-based conjoint analysis for these pricing decisions?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Submission

Make sure you have run all cells before submitting.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you have the zip file, upload the **entire** zip file to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 5)


Tests for q10 (points: 5)
```
>>> assert isinstance(convertible_worth_it_fun_segment, (bool, np.bool_))
```


Tests for q2 (points: 10)
```
>>> assert wtp_seat5 > 0 and wtp_auto > 0 and (wtp_convertible > 0), 'Some WTP values are not positive! Make sure you multiply the price coefficient by -1.'
```


Tests for q3 (points: 10)
```
>>> assert isinstance(thunderbolt_demand, np.ndarray), 'thunderbolt_demand should be a numpy array'
>>> assert len(thunderbolt_demand) == 25, 'thunderbolt_demand should have 25 elements (prices 26-50)'
```


Tests for q4 (points: 10)
```
>>> assert isinstance(thunderbolt_profit, np.ndarray), 'thunderbolt_profit should be a numpy array'
>>> assert len(thunderbolt_profit) == 25, 'thunderbolt_profit should have 25 elements'
```


Tests for q5 (points: 10)
```
>>> assert isinstance(thunderbolt_convertible_demand, np.ndarray), 'thunderbolt_convertible_demand should be a numpy array'
>>> assert len(thunderbolt_convertible_demand) == 25
```


Tests for q6 (points: 10)
```
>>> assert isinstance(thunderbolt_conv_profit, np.ndarray)
>>> assert 26 <= optimal_conv_price <= 50
```


Tests for q7 (points: 5)
```
>>> assert isinstance(convertible_worth_it, (bool, np.bool_)), 'convertible_worth_it should be a boolean'
```


Tests for q8 (points: 10)
```
>>> assert isinstance(coefs_fun, dict), 'coefs_fun should be a dictionary'
>>> assert 'convert' in coefs_fun and 'price' in coefs_fun
```


Tests for q9 (points: 10)
```
>>> assert 26 <= optimal_price_fun <= 50
>>> assert 26 <= optimal_conv_price_fun <= 50
```
