```python
# Initialize Otter
import otter
grader = otter.Notebook("hw0.ipynb")
```

# Homework 0: Setup and Understanding Otter Grader

```python
import polars as pl
import numpy as np
from plotnine import ggplot, aes, geom_line, labs, theme_minimal
```

**Question 1:** Create an array of numbers from 16 to 47 (inclusive) using NumPy.

```python
array = ...
```

```python
grader.check("q1")
```

**Question 2:** Load the data from the file `hw0_data.csv` into a Polars DataFrame and compute the mean and standard deviation of the `example` column.

Tips:
- To load a CSV file into a Polars DataFrame, you can use the `pl.read_csv()` function.
- To select a column (AKA Series) from a Polars DataFrame, you can use the syntax `df["column_name"]`.
- To compute the mean and standard deviation of a Polars Series, you can use the `.mean()` and `.std()` methods on the selected Series.

```python
df = ...
mean = ...
std = ...
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

**Question 3:** Plot the cumulative distribution function (CDF) of the `example` column from the DataFrame.

```python
# You can use the following function to plot the CDF. Read through it carefully to understand how it works.
def plot_cdf(dataframe, column_name):
    data = dataframe[column_name].to_numpy()
    sorted_data = np.sort(data)
    yvals = np.arange(1, len(sorted_data)+1) / float(len(sorted_data))
    
    plot_df = pl.DataFrame({'value': sorted_data, 'cumulative_probability': yvals})
    
    p = (
        ggplot(plot_df, aes(x='value', y='cumulative_probability'))
        + geom_line()
        + labs(x='Value', y='Cumulative Probability', title='Cumulative Distribution Function (CDF)')
        + theme_minimal()
    )
    return p
```

```python
# Plot the CDF of the 'example' column
...
```

<!-- END QUESTION -->

**Question 4 (multiple choice):** Based on the CDF plot, what percentage of values in the `example` column are less than or equal to 4?

Your answer can be one of the following options:
- A) Between 0% and 25%
- B) Between 25% and 50%
- C) Between 50% and 75%
- D) Between 75% and 100%

Write your answer by assigning a string variable `q4_answer` with the corresponding letter (e.g., `q4_answer = "A"`).

```python
q4_answer = ...
```

```python
grader.check("q4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you have the zip file, upload the **entire** zip file to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 25)
```
>>> len(array) == 32
True
```
```
>>> array[5] ** 3 == 9261
np.True_
```


Tests for q2 (points: 25)
```
>>> df.shape == (1000, 1)
True
```
```
>>> df['example'][673] == 4.867195284875214
True
```
```
>>> df['example'][41] == 2.6657239141808824
True
```


Tests for q4 (points: 25)
```
>>> q4_answer in ['A', 'B', 'C', 'D']
True
```
